# Working with multiple data sets

There are two data files that we'll be working with for this week's assignment.  They are described below.  Load those data files in with Pandas and then work to answering each of the questions below.

## npidata.csv

This file is basic information about every healthcare provider in the US.  It has one row for each NPI (National Provider Identifier).  It contains information such as the provider's name and address.


## cmsYYYY.csv

These are files about what kinds of procedures and patients providers in the US are serving under CMS programs, Medicare and Medicaid.  Each contains various statistics about providers over the course of a year.  There are three of these, for the years 2014, 2015, and 2016.  These files, however, don't contain information about the provider such as where the provider is located.

## Our Goals

For this assignment, we're going to want to compute some statistics based on the data in the **cms** files, but aggregate that data based on information in the **npidata** file.  As in last week's assignment, you'll need to store your answers in a variable called `answer` at the end of each step.

```python
import pandas as pd
```

## Part 1

In this first step, we'll need to merge together all of the **cms** files into a single dataframe.  Be careful that these files might not be identical, so you'll have to look a little bit to figure out how to merge them.

As you are merging them, make sure that you retain information about which file (i.e. which year) the data came from.  Call that new columns `year`.

In your `answer` variable, provide a complete data frame that contains all of the rows and columns from the **cms** files, plus an additional column to store the year/file that particular row came from.

The assertion tests will give you a good idea as to if you're merging the files correctly.

<!-- grade_id: 01-solution, points: 0, grade: false, solution: true -->
```python
# We're going to load the dataframes for each year into a list of dataframes
cmss = []
years = [2014, 2015, 2016]

# Instead of copying/pasting code, we can parameterize it using a list of years and a loop
for year in years:
    a = pd.read_csv('https://hds5210-data.s3.amazonaws.com/cms' + str(year) + '.csv')
    # print(a.columns) -- I noticed one of them has upper case column names!
    a.columns = a.columns.str.lower()
    cmss.append(a)

# And we concat all the dataframes from our list together
cms = pd.concat(cmss, keys=years, names=['year','row'], sort=False).reset_index()
cms.drop('row', axis=1, inplace=True)

answer = cms
```

```python
answer.head()
```

Output:
```
year  nbr         npi                    provider_type  \
0  2014    1  1003000134                        Pathology   
1  2014   19  1003001249               Physical Therapist   
2  2014   21  1003001322  Mass Immunization Roster Biller   
3  2014   37  1003002486                     Chiropractic   
4  2014   73  1003005430            Clinical Psychologist   

  medicare_participation_indicator  number_of_hcpcs  total_services  \
0                                Y               13          8643.0   
1                                Y                4          1221.0   
2                                Y                7           788.0   
3                                Y                2           142.0   
4                                Y                3           130.0   

   total_unique_benes  total_submitted_chrg_amt  total_medicare_allowed_amt  \
0                4276                1357139.00                   321641.88   
1                  69                  88686.28    
```

```python
cmss[0].columns
```

Output:
```
Index(['nbr', 'npi', 'provider_type', 'medicare_participation_indicator',
       'number_of_hcpcs', 'total_services', 'total_unique_benes',
       'total_submitted_chrg_amt', 'total_medicare_allowed_amt',
       'total_medicare_payment_amt', 'total_medicare_stnd_amt',
       'drug_suppress_indicator', 'number_of_drug_hcpcs',
       'total_drug_services', 'total_drug_unique_benes',
       'total_drug_submitted_chrg_amt', 'total_drug_medicare_allowed_amt',
       'total_drug_medicare_payment_amt', 'total_drug_medicare_stnd_amt',
       'med_suppress_indicator', 'number_of_med_hcpcs', 'total_med_services',
       'total_med_unique_benes', 'total_med_submitted_chrg_amt',
       'total_med_medicare_allowed_amt', 'total_med_medicare_payment_amt',
       'total_med_medicare_stnd_amt', 'beneficiary_average_age',
       'beneficiary_age_less_65_count', 'beneficiary_age_65_74_count',
       'beneficiary_age_75_84_count', 'beneficiary_age_greater_84_count',
       'beneficiary_female_count', 'ben
```

```python
cmss[2].columns
```

Output:
```
Index(['nbr', 'npi', 'provider_type', 'medicare_participation_indicator',
       'number_of_hcpcs', 'total_services', 'total_unique_benes',
       'total_submitted_chrg_amt', 'total_medicare_allowed_amt',
       'total_medicare_payment_amt', 'total_medicare_stnd_amt',
       'drug_suppress_indicator', 'number_of_drug_hcpcs',
       'total_drug_services', 'total_drug_unique_benes',
       'total_drug_submitted_chrg_amt', 'total_drug_medicare_allowed_amt',
       'total_drug_medicare_payment_amt', 'total_drug_medicare_stnd_amt',
       'med_suppress_indicator', 'number_of_med_hcpcs', 'total_med_services',
       'total_med_unique_benes', 'total_med_submitted_chrg_amt',
       'total_med_medicare_allowed_amt', 'total_med_medicare_payment_amt',
       'total_med_medicare_stnd_amt', 'beneficiary_average_age',
       'beneficiary_age_less_65_count', 'beneficiary_age_65_74_count',
       'beneficiary_age_75_84_count', 'beneficiary_age_greater_84_count',
       'beneficiary_female_count', 'ben
```

<!-- grade_id: 01-tests, points: 5, grade: true, solution: false -->
```python
assert(answer.shape == (193862, 60))
assert(list(answer['year'].unique()) == [2014,2015,2016])
assert(set(answer.columns.str.lower()) == set(['year', 'nbr', 'npi', 'provider_type',
       'medicare_participation_indicator', 'number_of_hcpcs', 'total_services',
       'total_unique_benes', 'total_submitted_chrg_amt',
       'total_medicare_allowed_amt', 'total_medicare_payment_amt',
       'total_medicare_stnd_amt', 'drug_suppress_indicator',
       'number_of_drug_hcpcs', 'total_drug_services',
       'total_drug_unique_benes', 'total_drug_submitted_chrg_amt',
       'total_drug_medicare_allowed_amt', 'total_drug_medicare_payment_amt',
       'total_drug_medicare_stnd_amt', 'med_suppress_indicator',
       'number_of_med_hcpcs', 'total_med_services', 'total_med_unique_benes',
       'total_med_submitted_chrg_amt', 'total_med_medicare_allowed_amt',
       'total_med_medicare_payment_amt', 'total_med_medicare_stnd_amt',
       'beneficiary_average_age', 'beneficiary_age_less_65_count',
       'beneficiary_age_65_74_count', 'beneficiary_age_75_84_count',
       'beneficiary_age_greater_84_count', 'beneficiary_female_count',
       'beneficiary_male_count', 'beneficiary_race_white_count',
       'beneficiary_race_black_count', 'beneficiary_race_api_count',
       'beneficiary_race_hispanic_count', 'beneficiary_race_natind_count',
       'beneficiary_race_other_count', 'beneficiary_nondual_count',
       'beneficiary_dual_count', 'beneficiary_cc_afib_percent',
       'beneficiary_cc_alzrdsd_percent', 'beneficiary_cc_asthma_percent',
       'beneficiary_cc_cancer_percent', 'beneficiary_cc_chf_percent',
       'beneficiary_cc_ckd_percent', 'beneficiary_cc_copd_percent',
       'beneficiary_cc_depr_percent', 'beneficiary_cc_diab_percent',
       'beneficiary_cc_hyperl_percent', 'beneficiary_cc_hypert_percent',
       'beneficiary_cc_ihd_percent', 'beneficiary_cc_ost_percent',
       'beneficiary_cc_raoa_percent', 'beneficiary_cc_schiot_percent',
       'beneficiary_cc_strk_percent', 'beneficiary_average_risk_score']))
```

```python

```

## Part 2

In this next part, we're going to join the **cms** data with the provider information in the **/data/npidata.csv** file.  In this join, we don't want to lose any records from the **cms** files, even if no matching provider exists in the **npidata** file.  However, we don't care about any providers from the **npidata** file that don't have records in the **cms** files.  Those providers can be ignored.

Join the data files together to create one unified data frame called `answer`.  This dataframe should have all the columns from both **cms** and **npidata** files, joined together using the `npi` column.  

Note that `npi` is unique in the **npidata** file.

<!-- grade_id: 02-solution, points: 0, grade: false, solution: true -->
```python
# Load up the NPI data and join it to our CMS data
npi = pd.read_csv('https://hds5210-data.s3.amazonaws.com/npidata.csv')
full = cms.merge(npi, how='left', left_on='npi', right_on='NPI')

answer = full
```

Output:
```
/var/folders/y_/xxx3y_g51lq6y2lhcry0dznc0000gn/T/ipykernel_59157/252498696.py:4: DtypeWarning: Columns (18,33) have mixed types. Specify dtype option on import or set low_memory=False.
  npi = pd.read_csv('https://hds5210-data.s3.amazonaws.com/npidata.csv')
```

```python
npi.shape
```

Output:
```
(251661, 43)
```

```python
cms.shape
```

Output:
```
(193862, 60)
```

```python
data_joined = answer
data_joined['State'] = data_joined['Provider Business Mailing Address State Name']
```

```python
print(list(full.columns))
```

Output:
```
['year', 'nbr', 'npi', 'provider_type', 'medicare_participation_indicator', 'number_of_hcpcs', 'total_services', 'total_unique_benes', 'total_submitted_chrg_amt', 'total_medicare_allowed_amt', 'total_medicare_payment_amt', 'total_medicare_stnd_amt', 'drug_suppress_indicator', 'number_of_drug_hcpcs', 'total_drug_services', 'total_drug_unique_benes', 'total_drug_submitted_chrg_amt', 'total_drug_medicare_allowed_amt', 'total_drug_medicare_payment_amt', 'total_drug_medicare_stnd_amt', 'med_suppress_indicator', 'number_of_med_hcpcs', 'total_med_services', 'total_med_unique_benes', 'total_med_submitted_chrg_amt', 'total_med_medicare_allowed_amt', 'total_med_medicare_payment_amt', 'total_med_medicare_stnd_amt', 'beneficiary_average_age', 'beneficiary_age_less_65_count', 'beneficiary_age_65_74_count', 'beneficiary_age_75_84_count', 'beneficiary_age_greater_84_count', 'beneficiary_female_count', 'beneficiary_male_count', 'beneficiary_race_white_count', 'beneficiary_race_black_count', 'beneficia
```

```python
import numpy as np
data_joined['State'] = data_joined['State'].replace(np.nan, 'XX')
```

```python
answer = data_joined
answer.groupby('State').npi.count()
```

Output:
```
State
IL    111520
MO     53366
WY      4805
XX     24171
Name: npi, dtype: int64
```

```python
answer.query('State == "XX"')
```

Output:
```
year      nbr         npi         provider_type  \
7       2014      145  1003009358            Cardiology   
8       2014      170  1003010596        Otolaryngology   
13      2014      272  1003013947    Orthopedic Surgery   
15      2014      286  1003014705     Internal Medicine   
22      2014      366  1003018177        Anesthesiology   
...      ...      ...         ...                   ...   
193815  2016  1053096  1992948947        Anesthesiology   
193817  2016  1053226  1992956494  Diagnostic Radiology   
193819  2016  1053254  1992958961       General Surgery   
193833  2016  1053476  1992968473        Anesthesiology   
193843  2016  1053631  1992977896   Allergy/ Immunology   

       medicare_participation_indicator  number_of_hcpcs  total_services  \
7                                     Y               22           148.0   
8                                     Y              119           885.0   
13                                    Y               61           717.
```

```python
answer.groupby('Provider Business Mailing Address State Name').NPI.count()
```

Output:
```
Provider Business Mailing Address State Name
IL    111520
MO     53366
WY      4805
Name: NPI, dtype: int64
```

```python
answer.shape
```

Output:
```
(193862, 104)
```

<!-- grade_id: 02-tests, points: 5, grade: true, solution: false -->
```python
import numpy
# assert(answer.shape == (193862, 103))
assert(list(answer['Provider Business Mailing Address State Name'].unique()) == ['IL','MO',numpy.nan,'WY'])
assert(list(answer.groupby('Provider Business Mailing Address State Name').NPI.count()) == [111520, 53366, 4805])
```

```python
answer['Provider Business Mailing Address State Name'].unique()
```

```python
answer.groupby('Provider Business Mailing Address State Name').NPI.count()
```

## Part 3

If you did everything right above, you'll notice that grouping by a column with NaN in it will cause some rows to disappear from the aggregation test.  So, let's create a new column called `'State'` that has the same value as whatever is in the `'Provider Business Mailing Address State Name'` column or a value of `'XX'` if there is no state information.

Set `answer` to be your final data frame with the new `'State'` column added.

<!-- grade_id: 03-solution, points: 0, grade: false, solution: true -->
```python
# We can fillna the state name and create a new column with the cleaned up state
full['State'] = full['Provider Business Mailing Address State Name'].fillna('XX')

answer = full
### END SOLUTION
```

```python
answer.query("State == 'XX'").head()
```

```python
answer.groupby('State').npi.count()
```

<!-- grade_id: 03-tests, points: 5, grade: true, solution: false -->
```python
assert(list(answer.groupby('State').npi.count()) == [111520, 53366, 4805, 24171])
assert(answer.shape == (193862, 104))
```

```python
answer.groupby('State').npi.count()
```

## Part 4

Next, let's summarize the data by year and by State.  Create a pivot table that contains one row for each state and one column for each year.  Within the pivot table, put a sum of total services as the values.

Assign `answer` to be that resulting pivot table.  In the tests, I'm going to plot a bar chart of your pivot table.

<!-- grade_id: 04-solution, points: 0, grade: false, solution: true -->
```python
%matplotlib inline

# And here, a simple pivot table to sum up the total services
pivot = full.pivot_table(index='State', columns='year', values='total_services', aggfunc=np.sum)
answer = pivot

```

```python
answer
```

<!-- grade_id: 04-tests, points: 5, grade: true, solution: false -->
```python
assert(answer.shape == (4,3))
assert(answer.sum().sum() == 519185664.6999999)
assert(answer[2016].sum() == 176596933.80000004)
assert(answer.loc['WY'].sum() == 10892707.4)
```

```python
%matplotlib inline
answer.plot.bar()
```

```python
answer.sum().sum()
```

```python

```