# Week 4 Exercises

See: _McKinney 3.1_ and [Python Documentation](https://docs.python.org/3/tutorial/datastructures.html) section 5 on data structures.

Modules 13 through 15 have all been about different built-in data types that can be used to create larger and more complex data structures.  We learned about tuples, lists, dictionaries, and how to access and process that data in new ways.

**At the begining of the semester, all of the workshop programming exercises will be structured a specific way to make it easier to verify for yourself that you're on the right track as well as easier for me to do a first pass on automated grading.  The structure of each question will require you to write a function using Python code. Don't worry that we haven't talked about functions yet. Just edit the code between** `### BEGIN SOLUTION` and `### END SOLUTION` **as shown in the example below.**


**WHAT I PROVIDE:**
```
def some_function(parameter1, parameter2):

   ### BEGIN SOLUTION
   x = -1
   ### END SOLUTION
   
   return x
```

**WHAT YOU SHOULD DO:** Just change the parameter names (if you feel you need to) and the calculations between `### BEGIN SOLUTION` and `### END SOLUTION`.  This is just a made up example.
```
def some_function(a, b):

   ### BEGIN SOLUTION
   temp = a + b
   x = temp / a * b
   ### END SOLUTION
   
   return x
```


---
---

**Below each programming exercise are some tests (`assertions`) that verify your code is working correctly.  If any assertions fail, you know that something isnt' right with your code, but having all assertions pass doesn't necessarily mean your code is perfect, yet. You should also create your own tests to make sure your code is correct.**

**For now, please don't change any function names**

### 20.1 Create a sorted list of names

Take the list of names below (some are first and last, others are first, middle, and last) and produce a final list of full names separated by commas.

```
sort_names([
  ['Paul', 'Boal'],
  ['Kermit','the','Frog'],
  ['Donald','Duck']])

'Paul Boal, Donald Duck, Kermit the Frog'
```

**Tip:** If you have a list of lists and want to sort it, Python can do that. It will simply sort using the first item in each of the lists inside your list.

```python
names = [['Paul', 'Boal'],
  ['Kermit', 'the', 'Frog'],
  ['Donald','Duck']]
```

```python
# Version 1: Create a new list with the last name at the front, sort it, and return just the joined name string

sortable_names = []
for name in names:
    sortable_names.append([name[-1], ' '.join(name)])
    # ['Last', 'First and Last']

sorted_names = []
for name in sorted(sortable_names):
    sorted_names.append(name[1])
    # Just ['First and Last']

print(sorted_names)
```

Output:
```
['Paul Boal', 'Donald Duck', 'Kermit the Frog']
```

```python
# Version 2: Using the sort key parameter

names.sort(key=lambda x: x[-1])
sorted_names = [' '.join(name) for name in names]

print(sorted_names)
```

Output:
```
['Paul Boal', 'Donald Duck', 'Kermit the Frog']
```

### 20.2 Convert a dictionary to a list

In this exercise, we're going to take a dictionary that has families and their names, and convert that into a list of names.  The dictionary is structured such that the key is the family name, and the value for that key is a list of first names.

```
{
  'Boal': ['Paul', 'Anny', 'Jim'],
  'Lester-Boal': ['Ellie', 'Ada', 'Teddy'],
  'Lester': ['Sarahlynn', 'Jessica', 'Grace', 'Carolynn']
}

[['Boal', 'Paul'],
 ['Boal', 'Anny'],
 ['Boal', 'Jim'],
 ['Lester-Boal', 'Elliie'],
 ['Lester-Boal', 'Ada'],
 ...
```

```python
families = {
  'Boal': ['Paul', 'Anny', 'Jim'],
  'Lester-Boal': ['Ellie', 'Ada', 'Teddy'],
  'Lester': ['Sarahlynn', 'Jessica', 'Grace', 'Carolynn']
}
```

```python
# The purpose here was to get you thinking about loops inside of loops
def dict_to_list(families):
    output = []

    # First, we create a loop to go through the dictionary
    for last, names in families.items():
        # For each dictionary item's value, we loop through the list of first names there
        for first in names:
            output.append([last, first])

    return output
```

```python
dict_to_list({'Duck': ['Hewey','Dewey','Louie'],
              'Mouse': ['Mickey','Minney']})
```

Output:
```
[['Duck', 'Hewey'],
 ['Duck', 'Dewey'],
 ['Duck', 'Louie'],
 ['Mouse', 'Mickey'],
 ['Mouse', 'Minney']]
```

### 20.3 Convert a list to a dictionary

Take the list of lists below and convert it into a dictionary where the entries in the dictionary are keyed off of whatever values are in column 0, and the dictionary values are the list of all the values that appear with that key from the input list.

<!-- grade_id: average-answer, points: 0, grade: false, solution: true -->
```python
def list_to_dict():

    names = [['Boal', 'Paul'],
             ['Duck', 'Donald'],
             ['Duck', 'Daisy'],
             ['Boal', 'Ada'],
             ['Boal', 'Teddy'],
             ['Westhus', 'Eric']]

    # To help you, I've provided a blank dictionary named `families`.
    families = {}

    # Your code should loop through `names` and do something to build
    # up `families` to be the output dictionary shown in the test
    # cells below.
    for last, first in names:
        families.setdefault(last, [])  # You could have also tested if last in names, but I like this one line shortcut
        families[last].append(first)   # The value for families[last] will be a list, so we append each first name to it

    ### BEGIN SOLUTION
    # Put your solution here
    ### END SOLUTION

    return families

```

```python
list_to_dict()
```

Output:
```
{'Boal': ['Paul', 'Ada', 'Teddy'],
 'Duck': ['Donald', 'Daisy'],
 'Westhus': ['Eric']}
```

<!-- grade_id: average-tests, points: 5, grade: true, solution: false -->
```python
answer = {
 'Boal':    ['Paul', 'Ada', 'Teddy'],
 'Duck':    ['Donald', 'Daisy'],
 'Westhus': ['Eric']
}

assert (list_to_dict() == answer)
```

### 20.4 Join using a dictionary

We have a list of patients, diagnosis, and length of stay.  We also have a dictionary that contains diagnosis and average length of stay.  Produce an output list that lists the patient and an indicator if the patient's stay was 'too long', 'too short', 'just right'

<!-- grade_id: cell-f3813e33810dc898, points: 0, grade: false, solution: true -->
```python
# The thing to realize here is that you don't have to loop through the dictionary for each patient.
# You can just access the avg_los value you want directly with avg_los.get(diagnosis)
def patient_los(avg_los, patients):

    conclusions = []

    # We loop through all the patients
    # This could have also been for name, diagnosis, actual_los in patients
    # Python automatically expands inner list items like that
    for pat in patients:
        name = pat[0]
        diagnosis = pat[1]
        actual_los = pat[2]
        # Here, we get the average / expected LOS
        target_los = avg_los.get(diagnosis)

        # Do our comparisons
        if actual_los == target_los:
            comment = "just right"
        elif actual_los < target_los:
            comment = "too short"
        else:
            comment = "too long"

        conclusions.append([name, comment])

    return conclusions
```

```python
avg_los = {
    "Hemolytic jaundice and perinatal jaundice" : 2,
    "Medical examination/evaluation" : 3.2,
    "Liveborn" : 3.2,
    "Trauma to perineum and vulva" : 2.1,
    "Normal pregnancy and/or delivery" : 2,
    "Umbilical cord complication" : 2.1,
    "Forceps delivery" : 2.2,
    "Administrative/social admission" : 4.2,
    "Prolonged pregnancy" : 2.4,
    "Other complications of pregnancy" : 2.5
}

patients = [
    ['Boal', 'Medical examination/evaluation', 1.1],
    ['Boal', 'Other complications of pregnancy', 3.3],
    ['Jones', 'Liveborn', 3.2],
    ['Ashbury', 'Forceps delivery', 2.0]
]

patient_los(avg_los, patients)
```

Output:
```
[['Boal', 'too short'],
 ['Boal', 'too long'],
 ['Jones', 'just right'],
 ['Ashbury', 'too short']]
```

<!-- grade_id: cell-2af0a520bb56ee21, points: 10, grade: true, solution: false -->
```python
conclusions = [['Boal', 'too short'],
    ['Boal', 'too long'],
    ['Jones', 'just right'],
    ['Ashbury', 'too short']]

assert(patient_los(avg_los, patients) == conclusions)
```

```python

```