# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 30.1 List of Allergies

In the /data directory on the Jupyter server, there is a file called `allergies.json` that contains a list of patient allergies.  It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json
ALLERGIES_FILE="allergies.json"
```

```python
def allergy_per_patient(json_file):
  """(dict) -> int
  This function does identifies the number of allergies per patient
  from a dictionary.

  >>> allergy_per_patient(ALLERGIES_FILE)
  {'Paul Boal': 1, 'Jason Argonaut': 3}
  """

  # First, we load up the data from the JSON formatted file into a Python object
  allergies = json.load(open(json_file))

  # We want the final output to be a dictionary of patient names,
  # so we start by initializing an empty dictionary.
  patient_allergies = {}

  # Now, we loop through all the entries in the input
  for allergy in allergies.get('entry'):

    # Each entry has a "patient" "resource" with a "display" name we can use for unique patient
    patient = allergy.get('resource').get('patient').get('display')

    if patient:
      # If there is a patient and it's one we've already seen, we'll keep counting
      if patient in patient_allergies:
        patient_allergies[patient] += 1
      # If it's a new patient we haven't seen before, we'll add it to the dictionary with a count of 1
      else:
        patient_allergies[patient] = 1

  # And return the final dictionary
  return patient_allergies
```

```python
# Process the file and show the output
allergy_per_patient('allergies.json')
```

Output:
```
{'Jason Argonaut': 3, 'Paul Boal': 1}
```

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
# Our allergy count can be very simple.
# Just load the file and count the entries.
def allergy_count(file):

    data = json.load(open(file))

    count = len(data["entry"])
    return count

```

```python
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count(ALLERGIES_FILE)) == int
assert allergy_count(ALLERGIES_FILE) == 4
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

```python
def patient_count(file):
    """(str) -> int
    Load the JSON data from the file and count how many unique patients there are,
    based on the patient's "display" name.
    """

    # Read the file and load the JSON formatted data into a Python object
    data = json.load(open(file))

    # Our output is a count of something unique, so we'll use a set here.
    # Remember a set is automatically a unique list of things.  So, it'll be easy to use.
    patients = set()

    # Loop through evetry entry and add the patient's display name to our set.
    # The way a set works, if the entry is already there, it won't error out.
    # It'll just silently move on.
    for entry in data["entry"]:
        patients.add(entry["resource"]["patient"]["display"])

    # In the end, we have a set of unique patient names and can just ask for the length of that set.
    return len(patients)
```

```python
patient_count(ALLERGIES_FILE)
```

Output:
```
2
```

```python
assert type(patient_count(ALLERGIES_FILE)) == int
assert patient_count(ALLERGIES_FILE) == 2
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


```python
def allergy_per_patient(file):

    data = json.load(open(file))

    patients = {}
    for entry in data["entry"]:
        name = entry["resource"]["patient"]["display"]
        allergy = entry["resource"]["substance"]["text"]
        # patients.setdefault(name, [])
        # patients[name].append(allergy)
        patients.setdefault(name, 0)
        patients[name] += 1

    return patients
```

```python
allergy_per_patient(ALLERGIES_FILE)
```

Output:
```
{'Jason Argonaut': 3, 'Paul Boal': 1}
```

```python
assert type(allergy_per_patient(ALLERGIES_FILE)) == dict
assert allergy_per_patient(ALLERGIES_FILE) == {'Paul Boal': 1, 'Jason Argonaut': 3}
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
def allergy_list(file):

    data = json.load(open(file))

    results = []
    for entry in data["entry"]:
        name = entry["resource"]["patient"]["display"]
        allergy = entry["resource"]["substance"]["text"]
        reaction = entry["resource"]["reaction"][0]["manifestation"][0]["text"]
        results.append([name, allergy, reaction])

    return results
```

```python
allergy_list(ALLERGIES_FILE)
```

Output:
```
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
assert allergy_list(ALLERGIES_FILE) == [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]

```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  You can solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
def allergy_reaction(file, patient, substance):

    data = json.load(open(file))

    results = []
    for entry in data["entry"]:
        name = entry["resource"]["patient"]["display"]
        allergy = entry["resource"]["substance"]["text"]
        reaction = entry["resource"]["reaction"][0]["manifestation"][0]["text"]
        if name == patient and allergy == substance:
            return reaction

    return None
```

```python
allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
```

Output:
```
'Hives'
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week06/week06_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

